<!--HEADER-->
[*Guia de aulas da disciplina de Modelagem Matemática*](https://github.com/rmsrosa/modelagem_matematica) *do* [*IM-UFRJ*](https://www.im.ufrj.br).

<!--BADGES-->
<a href="https://colab.research.google.com/github/rmsrosa/modelagem_matematica/blob/master/aulas/01.01-Teste1.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Google Colab" title="Open in Google Colab"></a>
&nbsp;<a href="https://mybinder.org/v2/gh/rmsrosa/modelagem_matematica/master?filepath=aulas/01.01-Teste1.ipynb"><img align="left" src="https://mybinder.org/badge.svg" alt="Binder" title="Open in binder"></a>
&nbsp;<a href="https://nbviewer.jupyter.org/github/rmsrosa/modelagem_matematica/blob/master/aulas/01.01-Teste1.ipynb"><img align="left" src="https://img.shields.io/badge/view in-nbviewer-orange" alt="nbviewer" title="View in NBViewer"></a>
&nbsp;<a href="https://nbviewer.jupyter.org/github/rmsrosa/modelagem_matematica/blob/master/aulas/slides/01.01-Teste1.slides.html"><img align="left" src="https://img.shields.io/badge/view-slides-darkgreen" alt="slides" title="View Slides"></a>
&nbsp;

<!--NAVIGATOR-->
[<- Princípios de Modelagem Matemática e primeiros exemplos](01.00.Aula-Modelagem.ipynb) | [Página Inicial](00.00-Pagina_Inicial.ipynb) 

---


# Teste

Defina, abaixo, uma função `fatorial` que recebe um inteiro não-negativo e retorne o seu fatorial, gerando erro caso o argumento não seja um inteiro positivo.

<!-- grade_id: cell-37099ea51404584d, points: 0, grade: false, solution: true -->
```python
def fatorial(n):
    ### Begin solution
    assert isinstance(n, int), "O argumento deve ser um inteiro."
    assert n >= 0, "O argumento deve ser um inteiro positivo."
    return 1 if n == 0 else n*fatorial(n-1)
    ### End solution
```

<!-- grade_id: cell-dc094241402e1474, points: 2, grade: true, solution: false -->
```python
for n, fat in ((0,1), (1,1), (2,2), (3,6), (4,24), (5,120), (10,3628800)):
    assert fatorial(n) == fat
```

```python
fatorial('2')
```

Output:
```
<ipython-input-15-8b1ebfcb6e76> in <module>
----> 1 fatorial('2')

<ipython-input-9-7e6a5a7f37ff> in fatorial(n)
      1 def fatorial(n):
      2     ### Begin solution
----> 3     assert isinstance(n, int), "O argumento deve ser um inteiro."
      4     assert n >= 0, "O argumento deve ser um inteiro positivo."
      5     return 1 if n == 0 else n*fatorial(n-1)

AssertionError: O argumento deve ser um inteiro.
```

```python

```

<!--NAVIGATOR-->

---
[<- Princípios de Modelagem Matemática e primeiros exemplos](01.00.Aula-Modelagem.ipynb) | [Página Inicial](00.00-Pagina_Inicial.ipynb) 